<!-- grade_id: ontdubbel_lijst, points: 0, grade: false, solution: true -->
```python
def ontdubbelList(lijst_met_dubbelen):
    """De input is een lijst. Return dezelfde lijst zonder dubbelen"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: ontdubbel_lijst_correct, points: 1, grade: true, solution: false -->
```python
assert ontdubbelList([1,1,2,3,4])==[1,2,3,4]
```

Output:
```
<ipython-input-3-f6472d557aac> in <module>
----> 1 assert ontdubbelList([1,1,2,3,4])==[1,2,3,4]

<ipython-input-2-87af9b4390e4> in ontdubbelList(lijst_met_dubbelen)
      2     """De input is een lijst. Return dezelfde lijst zonder dubbelen"""
      3     # YOUR CODE HERE
----> 4     raise NotImplementedError()

NotImplementedError:
```